<p style="font-family: Arial; font-size:3.75vw;color:purple; font-style:bold"><br>
Classification Exercise Notebook
</p><br>

# Exercise Notebook Instructions

### 1. Important: Only modify the cells which instruct you to modify them - leave "do not modify" cells alone.  

The code which tests your responses assumes you have run the startup/read-only code exactly.

### 2. Work through the notebook in order.

Some of the steps depend on previous, so you'll want to move through the notebook in order.

### 3. It is okay to use numpy libraries.

You may find some of these questions are fairly straightforward to answer using built-in numpy functions.  That's totally okay - part of the point of these exercises is to familiarize you with the commonly used numpy functions.

### 4. Seek help if stuck

If you get stuck, don't worry!  You can either review the videos/notebooks from this week, ask in the course forums, or look to the solutions for the correct answer.  BUT, be careful about looking to the solutions too quickly.  Struggling to get the right answer is an important part of the learning process.

```python
# DO NOT MODIFY

# import appropriate libraries
import pandas as pd
import numpy as np
from sklearn.model_selection import train_test_split
from sklearn.metrics import accuracy_score
```

```python
# DO NOT MODIFY

# we will use the daily weather dataset of this exercise.

def get_data():
    return pd.read_csv('../Week-7-MachineLearning/weather/daily_weather.csv')

df = get_data()
```

```python
df.head()
```

Output:
```
number  air_pressure_9am  air_temp_9am  avg_wind_direction_9am  \
0       0        918.060000     74.822000              271.100000   
1       1        917.347688     71.403843              101.935179   
2       2        923.040000     60.638000               51.000000   
3       3        920.502751     70.138895              198.832133   
4       4        921.160000     44.294000              277.800000   

   avg_wind_speed_9am  max_wind_direction_9am  max_wind_speed_9am  \
0            2.080354              295.400000            2.863283   
1            2.443009              140.471548            3.533324   
2           17.067852               63.700000           22.100967   
3            4.337363              211.203341            5.190045   
4            1.856660              136.500000            2.863283   

   rain_accumulation_9am  rain_duration_9am  relative_humidity_9am  \
0                    0.0                0.0              42.420000   
1                    0.0         
```

<p style="font-family: Arial; font-size:2.75vw;color:purple; font-style:bold"><br>

Exercise 1: Detect if NULLs Exist in the Data<br><br></p>


In the cell below, modify the function to RETURN True if there exists a NULL value in data given as input, otherwise False

<!-- grade_id: c555, points: 0, grade: false, solution: true -->
```python
# modify this cell

def detect_nulls(matx):
    ### BEGIN SOLUTION
    return matx.isnull().any().any()
    ### END SOLUTION
```

<!-- grade_id: ex15, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY
ans = True

try: assert np.alltrue(detect_nulls(df) == ans)
except AssertionError as e: print("Try again, your output did not match the expected answer above")
```

<p style="font-family: Arial; font-size:2.75vw;color:purple; font-style:bold"><br>

Exercise 2: Data Cleaning<br><br></p>


In the cell below, modify the function to remove every row that has atleast one NULL,
and RETURN a new data frame.

<!-- grade_id: 5551, points: 0, grade: false, solution: true -->
```python
# modify this cell

def clean_data(matx):
    ### BEGIN SOLUTION
    x = matx.dropna()
    return x
    ### END SOLUTION
```

<!-- grade_id: ex151, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY

try: 
    df2=clean_data(df)
    assert np.alltrue(df2.shape[0] == 1064)
except AssertionError as e: print("Try again")
```

```python
df2.describe()
```

Output:
```
number  air_pressure_9am  air_temp_9am  avg_wind_direction_9am  \
count  1064.000000       1064.000000   1064.000000             1064.000000   
mean    545.001880        918.903180     65.022609              142.306756   
std     316.796111          3.179040     11.168033               69.149472   
min       0.000000        907.990000     36.752000               15.500000   
25%     269.750000        916.595376     57.398000               65.979244   
50%     543.500000        918.942281     65.778479              165.937461   
75%     819.250000        921.169054     73.530872              191.100000   
max    1094.000000        929.320000     98.906000              343.400000   

       avg_wind_speed_9am  max_wind_direction_9am  max_wind_speed_9am  \
count         1064.000000             1064.000000         1064.000000   
mean             5.485793              148.480424            6.999714   
std              4.534427               67.154911            5.590790   
min              
```

<p style="font-family: Arial; font-size:2.75vw;color:purple; font-style:bold"><br>

Exercise 3: Convert to a Classification Task<br><br></p>

In the cell below, modify the function that takes three inputs: a dataframe, a column name C, and a threshold value X.
The function has to create a new column called 'target' which is 1
when given column C is > X, otherwise 0. Return a new dataframe with this new column, and **keep all the 
previous columns.**

<!-- grade_id: c5552, points: 0, grade: false, solution: true -->
```python
# modify this cell

def classification_task(df, C, X):
    ### BEGIN SOLUTION
    df_copy = df.copy()
    df_copy['target'] = (df[C] > X)*1
    return df_copy
    ### END SOLUTION
```

```python
#mia soluzione
def classification_task(df, C, X):
        df['target'] = (df.C > X)*1
        return df
```

<!-- grade_id: ex152, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY

ans = 558

try: 
    c = classification_task(df2, 'air_temp_9am', 65.0)
    assert np.alltrue(c['target'].sum() == ans)
except AssertionError as e: print("Try again, your solution did not produce the expected output above")
```

```python
del c['number']
```

```python
def get_ready(df):
    from sklearn.model_selection import train_test_split
    # declare features to use
    features = ['air_pressure_9am','avg_wind_direction_9am','avg_wind_speed_9am','max_wind_direction_9am']
    # create copy
    df_copy = df.copy()
    df_copy['target'] = (df_copy['air_temp_9am'] > 65.0)*1
    # separate features from target
    X=df_copy[features]
    y=df_copy['target']
    # perform split and return
    return train_test_split(X, y, test_size=0.33, random_state=983)
```

```python
X_train, X_test, y_train, y_test = get_ready(pd.read_csv('../Week-7-MachineLearning/weather/daily_weather.csv').dropna())
```

<p style="font-family: Arial; font-size:2.75vw;color:purple; font-style:bold"><br>

Exercise 4: Create a Classifier<br><br></p>

In the cell below, modify the function to take X_train, y_train only and RETURN a classifier
for predicting the y_train based on columns in X_train. You can pick any classifier model.

The function should RETURN a trained model. We will test your classifier on X_test and y_test

<!-- grade_id: c5553, points: 0, grade: false, solution: true -->
```python
# modify this cell

def train_classifier(X_train, y_train):
    ### BEGIN SOLUTION
    
    #from sklearn.tree import DecisionTreeClassifier
    #_classifier = DecisionTreeClassifier(max_leaf_nodes = 30, random_state = 0)
    
    from sklearn import svm
    _classifier = svm.SVC(gamma = 'auto')
    
    #from sklearn.linear_model import SGDClassifier
    #_classifier = SGDClassifier(max_iter = 1000)
    
    
    _classifier.fit(X_train, y_train)
    
    return _classifier

    ### END SOLUTION
```

<!-- grade_id: ex153, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY

min_score = 0.63

try: 
    v = train_classifier(X_train, y_train)
    y_pred = v.predict(X_test)
    assert np.alltrue(accuracy_score(y_test, y_pred) > min_score)
except AssertionError as e: print("Keep trying - can you get an accuracy_score > %f" % min_score)
```

```python
accuracy_score(y_true = y_test, y_pred = v.predict(X_test))
```

Output:
```
0.6306818181818182
```

```python

```